```python
# Pandas Practice Problems
```

```python
%load_ext pycodestyle_magic
%flake8_on --ignore E266 --max_line_length 99
```

Output:
```
The pycodestyle_magic extension is already loaded. To reload it, use:
  %reload_ext pycodestyle_magic
```

```python
from urllib import request
from os import path
import pandas as pd
import numpy as np

url = "https://archive.ics.uci.edu/ml/machine-learning-databases/forest-fires/"
fn = "forestfires.csv"
if not path.exists(fn):
    data = request.urlopen(url + fn).read()
    with open(fn, 'wb') as fhand:
        fhand.write(data)

ff = pd.read_csv("forestfires.csv")
ff

# print(ff.info())
```

Output:
```
X  Y month  day  FFMC    DMC     DC   ISI  temp  RH  wind  rain   area
0    7  5   mar  fri  86.2   26.2   94.3   5.1   8.2  51   6.7   0.0   0.00
1    7  4   oct  tue  90.6   35.4  669.1   6.7  18.0  33   0.9   0.0   0.00
2    7  4   oct  sat  90.6   43.7  686.9   6.7  14.6  33   1.3   0.0   0.00
3    8  6   mar  fri  91.7   33.3   77.5   9.0   8.3  97   4.0   0.2   0.00
4    8  6   mar  sun  89.3   51.3  102.2   9.6  11.4  99   1.8   0.0   0.00
..  .. ..   ...  ...   ...    ...    ...   ...   ...  ..   ...   ...    ...
512  4  3   aug  sun  81.6   56.7  665.6   1.9  27.8  32   2.7   0.0   6.44
513  2  4   aug  sun  81.6   56.7  665.6   1.9  21.9  71   5.8   0.0  54.29
514  7  4   aug  sun  81.6   56.7  665.6   1.9  21.2  70   6.7   0.0  11.16
515  1  4   aug  sat  94.4  146.0  614.7  11.3  25.6  42   4.0   0.0   0.00
516  6  3   nov  tue  79.5    3.0  106.7   1.1  11.8  31   4.5   0.0   0.00

[517 rows x 13 columns]
```

### Problem 1: Filtering and computing a mean
Write a function that returns the average rainfall on days with wind speed above 6 mph.

```python
ff.info()


```

Output:
```
<class 'pandas.core.frame.DataFrame'>
RangeIndex: 517 entries, 0 to 516
Data columns (total 13 columns):
 #   Column  Non-Null Count  Dtype  
---  ------  --------------  -----  
 0   X       517 non-null    int64  
 1   Y       517 non-null    int64  
 2   month   517 non-null    object 
 3   day     517 non-null    object 
 4   FFMC    517 non-null    float64
 5   DMC     517 non-null    float64
 6   DC      517 non-null    float64
 7   ISI     517 non-null    float64
 8   temp    517 non-null    float64
 9   RH      517 non-null    int64  
 10  wind    517 non-null    float64
 11  rain    517 non-null    float64
 12  area    517 non-null    float64
dtypes: float64(8), int64(3), object(2)
memory usage: 52.6+ KB
```

<!-- grade_id: cell-a302cac2faa1187d, points: 0, grade: false, solution: true -->
```python
def windy_rain():
    windy_days = ff[ff['wind']>6]
    print(windy_days)
    return windy_days['rain'].mean()
    # YOUR CODE HERE
```

Output:
```
Error in callback <bound method VarWatcher.auto_run_flake8 of <pycodestyle_magic.VarWatcher object at 0x00000169481417F0>> (for post_run_cell), with arguments args (<ExecutionResult object at 1694c2ca4a0, execution_count=18 error_before_exec=None error_in_exec=None info=<ExecutionInfo object at 1694cd4bb50, raw_cell="def windy_rain():
    windy_days = ff[ff['wind']>6.." transformed_cell="def windy_rain():
    windy_days = ff[ff['wind']>6.." store_history=True silent=False shell_futures=True cell_id=b7b1e895-662a-490b-b067-abc4101972ea> result=None>,),kwargs {}:
File C:\Anaconda\Lib\site-packages\flake8\formatting\base.py:194, in BaseFormatter.write(self, line, source)
    181 """Write the line either to the output file or stdout.
    182 
    183 This handles deciding whether to write to a file or print to standard
   (...)    191     line of output.
    192 """
    193 if line:
--> 194     self._write(line)
    195 if source:
    196     self._write(source)

File C:\Anaconda\Lib\site-
```

<!-- grade_id: cell-0bcdc1e3800f6a24, points: 1, grade: true, solution: false -->
```python
assert abs(windy_rain() - 0.0412698) < 1e-6
```

Output:
```
X  Y month  day  FFMC    DMC     DC   ISI  temp  RH  wind  rain   area
0    7  5   mar  fri  86.2   26.2   94.3   5.1   8.2  51   6.7   0.0   0.00
10   7  5   sep  sat  92.5   88.0  698.6   7.1  17.8  51   7.2   0.0   0.00
12   6  5   aug  fri  63.5   70.8  665.3   0.8  17.0  72   6.7   0.0   0.00
31   6  3   sep  mon  88.6   91.8  709.9   7.1  11.2  78   7.6   0.0   0.00
41   4  4   aug  sat  90.2   96.9  624.2   8.9  18.4  42   6.7   0.0   0.00
..  .. ..   ...  ...   ...    ...    ...   ...   ...  ..   ...   ...    ...
486  5  4   aug  tue  95.1  141.3  605.8  17.7  24.1  43   6.3   0.0   2.00
488  4  4   aug  tue  95.1  141.3  605.8  17.7  19.4  71   7.6   0.0  46.70
506  1  2   aug  fri  91.0  166.9  752.6   7.1  18.5  73   8.5   0.0   0.00
509  5  4   aug  fri  91.0  166.9  752.6   7.1  21.1  71   7.6   1.4   2.17
514  7  4   aug  sun  81.6   56.7  665.6   1.9  21.2  70   6.7   0.0  11.16

[63 rows x 13 columns]
```

### Problem 2: Filtering and computing a sum
Write a function that take a month as an argument and returns the total rainfall during the month.

<!-- grade_id: cell-a21234336020693a, points: 0, grade: false, solution: true -->
```python
def monthly_rain(month):
    # YOUR CODE HERE
```

<!-- grade_id: cell-db413049a2a4d48e, points: 1, grade: true, solution: false -->
```python
assert abs(monthly_rain("aug") - 10.8) < 1e-6
assert abs(monthly_rain("mar") - 0.2) < 1e-6
```

### Problem 3: Loading Pandas data frames
Download the white wine dataset from [https://archive.ics.uci.edu/ml/machine-learning-databases/wine-quality/winequality-white.csv](https://archive.ics.uci.edu/ml/machine-learning-databases/wine-quality/winequality-white.csv), and load it into a Pandas dataframe called `df`.

<!-- grade_id: cell-a24f7cb6ee8b606c, points: 0, grade: false, solution: true -->
```python
from urllib import request
from os import path
import pandas as pd

url = "https://archive.ics.uci.edu/ml/machine-learning-databases/wine-quality/"
fn = "winequality-white.csv"
# YOUR CODE HERE
```

<!-- grade_id: cell-d566adf4a0305bf4, points: 1, grade: true, solution: false -->
```python
assert abs(df["citric acid"].mean() - 0.3341915) < 1e-8
```

### Problem 4: pipe
Write a function that takes a user-defined function as an argument.
It uses `pipe` to apply the function to each wine pH, and returns the mean of the values returned by the user-defined function.

<!-- grade_id: cell-3c158319d4eb0923, points: 0, grade: false, solution: true -->
```python
def mean_ph(user_fun):
    # YOUR CODE HERE
```

<!-- grade_id: cell-6d36ddd9e5bea7d4, points: 1, grade: true, solution: false -->
```python
assert abs(mean_ph(lambda x: x**2) - (df.pH ** 2).mean()) < 1e-8
```

```python

```